---
title: 'Using Parts of a Date'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Using Parts of a Date

```yaml
type: VideoExercise
key: 984d436696
xp: 50
```

`@projector_key`
b74e45a133085fdd361db173130ce4d1
